---
title: 'TIPOS DE DATOS. VARIABLES'
description: "Tipos de datos en Python. \nVariables"
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

Los tipos de datos básicos en Python son int (nº entero), float (nº real), string (cadena de caracteres) y bool (booleano).
Las variables en Python se crean sin 'tipar', escribiendo el nombre en minúsculas y asignando el valor después del signo =

`@instructions`
Crea las variables mi_int, mi_float, mi_string y mi_bool. Asígnales los valores 7, 3.12, "hola" y True respectivamente

`@hint`
Recuerda que la sintaxis para crera una variable es: nombre_variable=valor

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
mi_int = 7
```

`@solution`
```{python}
mi_int = 7 
mi_float = 3.12
mi_string = "hola"
mi_bool = True 
```

`@sct`
```{python}

```
